Strip cropping shows promising increases in ground beetle community diversity compared to monocultures

Author response:

The following is the authors’ response to the previous reviews

Reviewer #3 (Public Review)

Summary:

In this paper the authors examined the effects of strip cropping, a relatively new agricultural technique of alternating crops in small strips of several meters wide, on ground beetle diversity. The results show an increase in species diversity (i.e. abundance and species richness) of the ground beetle communities compared to monocultures.

Strengths:

The article is well written; it has an easily readable tone of voice without too much jargon or overly complicated sentence structure. Moreover, as far as reviewing the models in depth without raw data and R scripts allows, the statistical work done by the authors looks good. They have well thought out how to handle heterogenous, unbalanced and taxonomically unspecific yet spatially and temporarily correlated field data. The models applied and the model checks performed are appropriate for the data at hand. Combining RDA and PCA axes together is a nice touch. Moreover, after the first round of reviews, the authors have done a great job at rewriting the paper to make it less overstated, more relevant to the data at hand and more solid in the findings. Many of the weaknesses noted in the first review have been dealt with. The overall structure of the paper is good, with a clear introduction, hypotheses, results section and discussion.

We are grateful for this positive feedback. We are glad that our extensive revision after extensive review from three reviewers has paid off in addressing earlier weakness of our manuscript.

Weaknesses:

The weaknesses that remain are mainly due to a difficult dataset and choices that could have stressed certain aspects more, like the relationship between strip cropping and intercropping. The mechanistic understanding of strip cropping is what is at stake here. Does strip cropping behave similar to intercropping, a technique which has been proven to be beneficial to biodiversity because of added effects due to increased resource efficiency and greater plant species richness.

Unfortunately, the authors do not go into this in the introduction or otherwise and simply state that they consider strip cropping a form of intercropping.

We agree with the reviewer that a mechanistic understanding on how intercropping and strip cropping differ would be very interesting. However, we also feel that this topic is somewhat beyond the scope of the current manuscript. We are already planning work to elucidate mechanisms that may explain the pest and suppressive effects of strip cropping.

I also do not like the exclusive focus on percentages, as these are dimensionless. I think more could have been done to show underlying structure in the data, even after rarefaction.

While we generally agree with this point raised by the reviewer, for our heterogeneous dataset it was difficult to come up with meaningful units with dimensions. Therefore, we believe that percentages are the most suitable approach to present readers a fair comparison of the treatments.

A further weakness is a limited embedding into the larger scientific discourses other than providing references. But this may be a matter of style and/or taste

We believe our manuscript to be well-embedded within the relevant scientific discourse, but as indicated by reviewer 3 this might indeed be a matter of style/taste. Without exact examples it is difficult for us to judge this point.

Reviewer #3 (Recommendations for the authors):

Suggestion for title: "Strip cropping shows promising preliminary increases in ground beetle community diversity compared to monocultures"

We agree that the title could indeed be nuanced. We incorporated the suggested title, except for the word “preliminary”, as we felt that this is slightly misplaced for a 4-year study conducted at 4 locations.

line 26: the word previous may be confusing to readers, as it suggests previous research on beetles or insects. I think it would be better to use for instance "related" or "productivity focused research"

We agree that this wording might be confusing, and changed it to “other studies showed”.

Line 84-85: this is vague. can you make explicit what you are trying to answer here?

We made “biodiversity metric changes” more explicit, and changed the sentence accordingly.

Line 88-89: I think this would fit better with the first question in line 83-84, so I suggest placing it upwards. Also, I think you mean abundant instead of common. Common suggests commonness in the entire population. Abundant suggests found often in this study. While these definitions may very much overlap, they are distinctly different.

We have moved this sentence up and changed “common” to “abundant”. To make the result section more in line with this section, we also moved the section on the relationship between crop configuration and abundant genera up.

Line 146: defining rareness of species should be in the methods section. Also "following" would be better than "according"

We now added a sentence on how we examine habitat preferences and rarity in the methods section (line 316-317). We also changed “according to” to “following”.

Line 291: it is called being "flush" with the soil surface. This expression is not much used by non-native speakers, but is regularly encountered in studies on pitfalls, so the authors could decide to change the sentence using the proper English vernacular.

Suggestion incorporated.

Line 322-327, this method could do with a reference

This method is a relatively standard calculation to calculate relative changes and to center variation around zero. Nevertheless, we added a reference to a paper that used the same method.

Line: 333-335. I would still like to see a reference for this method.

This methodology has not been described in literature to the best of our knowledge. As we compared two crops within strip cropping with their respective monoculture references, we compare one strip cropping field with two monocultural fields. Here we took a conservative approach by comparing the strip crop field with the monoculture with the highest richness and activity density, to see if strip cropped fields outperformed monocultures with diverse ground beetle communities.

Line 364-366. references?

We have added references for these R packages.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Global biodiversity is declining at an unprecedented rate, with agriculture as one of the major drivers. There is mounting evidence that intercropping can increase insect biodiversity while maintaining or increasing yield. Yet, intercropping is often considered impractical for mechanized farming systems. Strip cropping is a type of intercropping that is compatible with standard farm machinery and has been pioneered by Dutch farmers since 2014. Here, we present ground beetle data from four organically managed experimental farms across four years. Ground beetles are sensitive to changes in habitats and disturbances, and hold keystone positions in agroecosystem food webs. We show that strip cropping systems can enhance ground beetle biodiversity, while previous research shows that these increases have been achieved without incurring major yield loss. Strip cropped fields had on average 15% more ground beetle species and 30% more individuals than monocultural fields. The higher ground beetle richness in strip crops was explained by the merger of crop-related ground beetle communities, rather than by ground beetle species unique to strip cropping systems. The increase in field-level beetle species richness in organic agriculture through strip cropping approached increases found for other readily deployed biodiversity conservation methods, like shifting from conventional to organic agriculture (+19% - +23%). This indicates that strip cropping is a potentially useful tool supporting ground beetle biodiversity in agricultural fields without compromising food production.

Revised version: Global biodiversity is declining at an unprecedented rate, with agriculture as one of the major drivers. There is mounting evidence that intercropping can increase insect biodiversity while maintaining or increasing yield. Yet, intercropping is often considered impractical for mechanized farming systems. Strip cropping is a type of intercropping that is compatible with standard farm machinery and has been pioneered by Dutch farmers since 2014. Here, we present ground beetle data from four organically managed experimental farms across four years. Ground beetles are sensitive to changes in habitats and disturbances, and hold keystone positions in agroecosystem food webs. We show that strip cropping systems can enhance ground beetle biodiversity, while other studies showed that these increases have been achieved without incurring major yield loss. Strip cropped fields had on average 15% more ground beetle species and 30% more individuals than monocultural fields. The higher ground beetle richness in strip crops was explained by the merger of crop-related ground beetle communities, rather than by ground beetle species unique to strip cropping systems. The increase in field-level beetle species richness in organic agriculture through strip cropping approached increases found for other readily deployed biodiversity conservation methods, like shifting from conventional to organic agriculture (+19% - +23%). This indicates that strip cropping is a potentially useful tool supporting ground beetle biodiversity in agricultural fields without compromising food production.

Introduction

Earlier version: Insects account for 80% of the animal species in the world and, therefore, the recently reported unprecedented rate of insect decline is cause for alarm about the state of biodiversity on Earth (Dirzo et al., 2014; Hallmann et al., 2017; Van Klink et al., 2020). An array of biodiversity metrics provide strong evidence for declines of especially terrestrial insects across all continents (Van Klink et al., 2020). This includes declines in abundance of both common and rare species (Hallmann et al., 2017; Pilotto et al., 2020; Seibold et al., 2019; van Klink et al., 2024) and in total species richness and changes in assembly composition (Blowes et al., 2022; Wagner et al., 2021). Insects are essential for crop production through their role in decomposition, pollination, pest control and sustaining food webs. Therefore, erosion of insect communities can have potentially devastating effects on ecosystem functioning, provision of ecosystem services and ultimately on human civilization (Dirzo et al., 2014). The main drivers of insect biodiversity decline are habitat loss due to conversion to agriculture, pollution, invasive species and climate change (Díaz et al., 2019; Müller et al., 2023; Wagner et al., 2021). Strategies for biodiversity conservation in conjunction with adequate food production require understanding of how biodiversity responds to agricultural management (Cozim-Melges et al., 2024; Mupepele et al., 2021; Saunders, 2020). Ideally, sustainable agricultural practices retain yield and enhance biodiversity, preventing the need to convert natural habitats to agricultural land to maintain food production (Tscharntke et al., 2021).

Revised version: Insects account for 80% of the animal species in the world and, therefore, the recently reported unprecedented rate of insect decline is cause for alarm about the state of biodiversity on Earth (Dirzo et al., 2014; Hallmann et al., 2017; Van Klink et al., 2020). An array of biodiversity metrics provide strong evidence for declines of especially terrestrial insects across all continents (Van Klink et al., 2020). This includes declines in abundance of both common and rare species (Hallmann et al., 2017; Pilotto et al., 2020; Seibold et al., 2019; van Klink et al., 2024) and in total species richness and changes in assembly composition (Blowes et al., 2022; Wagner et al., 2021). Insects are essential for crop production through their role in decomposition, pollination, pest control and sustaining food webs. Therefore, erosion of insect communities can have potentially devastating effects on ecosystem functioning, provision of ecosystem services and ultimately on human civilization (Dirzo et al., 2014). The main drivers of insect biodiversity decline are habitat loss due to conversion to agriculture, pollution, invasive species and climate change (Díaz et al., 2019; Müller et al., 2024; Wagner et al., 2021). Strategies for biodiversity conservation in conjunction with adequate food production require understanding of how biodiversity responds to agricultural management (Cozim-Melges et al., 2024; Mupepele et al., 2021; Saunders, 2020). Ideally, sustainable agricultural practices retain yield and enhance biodiversity, preventing the need to convert natural habitats to agricultural land to maintain food production (Tscharntke et al., 2021).

Increasing crop heterogeneity can facilitate biodiversity conservation in highly productive agricultural landscapes without compromising yield (Martin-Guay et al., 2018; Sirami et al., 2019). Crop diversification can enhance niche complementarity by creating heterogeneous habitats and increasing availability and diversity of resources (Lichtenberg et al., 2017; Tamburini et al., 2020). A promising crop diversification strategy is strip cropping, where crops are grown in alternating strips, wide enough for using standard agricultural machines yet narrow enough to facilitate ecological interactions among crops (Ditzler et al., 2021). Crops that are grown in strips may benefit from increased resource use efficiency, and the suppression of pests and diseases (Croijmans et al., 2024; Karssemeijer et al., 2024; Rakotomalala et al., 2023) without major yield compromises (Campanelli et al., 2023; Juventia and van Apeldoorn, 2024; Van Oort et al., 2020). Growing multiple crops on a field may foster a larger diversity of organisms than monocultures through greater plant species richness that cascades into richer herbivore and predator communities (Crutsinger et al., 2006; Cuperus et al., 2023). Moreover, the expected increase in available and potentially complementary niches within the agricultural field due to higher spatial diversity in crops can result in admixture of communities related to individual crops (Hummel et al., 2012), or the creation of completely new communities by enhanced richness of agriculture-related species, and/or the occurrence of species rarely found in agricultural fields (Rischen et al., 2021). So far, it is not well understood if and how insect communities respond to strip cropping across distinct crops.

Earlier version: Here, we present data of four years of pitfall trapping of ground beetles at several moments during the growing season in 14 crops at four organic experimental farms across the Netherlands where strip cropping was compared to monocultures. Ground beetle communities are sensitive to changes in farming practices and are frequently used to examine agricultural sustainability (Holland and Luff, 2000; Makwela et al., 2023; Turin, 2022). Furthermore, ground beetle species are important for maintaining ecological functions as they comprise scavengers and predators of (weed) seeds, detritivores (e.g., collembolas and earthworms) and herbivores (e.g., aphids and caterpillars). We first examine whether strip cropping fields have greater ground beetle richness and activity density than monocultural fields. Next, we present biodiversity metric changes for individual crops to assess the benefits of strip cropping for ground beetle biodiversity. We then evaluate whether ground beetle community changes are caused by admixture of communities, whether these assemblages promote species associated with agricultural or with natural ecosystems, and whether they contain rare species. Lastly, we test for 12 common ground beetle genera whether their activity density is higher in strip cropping compared to monocultures.

Revised version: Here, we present data of four years of pitfall trapping of ground beetles at several moments during the growing season in 14 crops at four organic experimental farms across the Netherlands where strip cropping was compared to monocultures. Ground beetle communities are sensitive to changes in farming practices and are frequently used to examine agricultural sustainability (Holland and Luff, 2000; Makwela et al., 2023; Turin, 2022). Furthermore, ground beetle species are important for maintaining ecological functions as they comprise scavengers and predators of (weed) seeds, detritivores (e.g., collembolas and earthworms) and herbivores (e.g., aphids and caterpillars). We first examine whether strip cropping fields have greater ground beetle activity density, species richness, evenness and diversity than monocultural fields. We also test for twelve abundant ground beetle genera whether their activity density is higher in strip cropping than in monocultures. Lastly, we evaluate whether ground beetle community changes are caused by admixture of communities, whether these assemblages promote species associated with agricultural or with natural ecosystems, and whether they contain rare species.

Results

A total of 48,108 ground beetles belonging to 71 species were caught using pitfall traps over four years at four different organically managed experimental farms in The Netherlands: 40,153 at Almere; 3,777 at Lelystad; 1,126 at Valthermond; and 3,052 at Wageningen (Table S1, S2).

Strip cropping enhances ground beetle richness

Strip cropping fields had on average 15% higher ground beetle taxonomic richness than monoculture fields after rarefaction to the number of samples of the least-sampled crop configuration (β0 = 0.151, SE = 0.044, p < 0.001; Fig. 1c,d). However, strip cropping fields did not harbour more species than monocultural fields with the highest ground beetle richness (β0 = - 0.008, SE = 0.037, p = 0.821; Fig. 1d). The difference in field-level taxonomic richness could not be explained by an increase in the number of ground beetle species per crop in strip cropping compared to monocultures. At crop-level, the 5% increase in ground beetle taxonomic richness in strip cropping was not statistically significant (Fig. 1f). Similarly, crop-level absolute evenness, inverse Simpson index and Shannon entropy did not differ significantly among crop configurations (Fig. 1h-j, S1c,d, S2). The effect of crop configuration on crop-level taxonomic richness was variable and was not associated with location or crop species. The effect ranged from 56% more species in potato in monoculture at Wageningen in 2022, to 136% more species in barley in strip cropping at Wageningen in 2020 (Fig. 1k).

Effect of crop configuration (monoculture versus strip cropping) on ground beetle biodiversity.

(a) Location of experimental sites in the Netherlands. (b) Field set-up of the two crop configurations: monoculture and strip cropping. At Lelystad and Wageningen, strip cropping consisted of 3-m wide crop strips of two crops (pairs), and multiple crop pairs were assessed. At Almere and Valthermond, strip cropping consisted of 6-m wide crop strips of eight crops combined (Fig. S10). (c) Sample-based species accumulation curves of all year series from monocultures (brown) and strip cropping (green), in Almere from 2021 and 2022. This was the only location where an equal number of samples were taken in the monocultures and in strip cropping on a similar area. Ground beetle species include Poecilus cupreus (left) and Pterostichus melanarius (right). Photo credit: Ortwin Bleich, retrieved from: www.eurocarabidae.de. (d-e) Overall relative change in field-level ground beetle (d) taxonomic richness and (e) activity density. Positive values indicate higher richness or activity density in strip cropping, negative values in monocultures. (f-j) Overall relative effect of crop configuration on ground beetle (f) taxonomic richness, (g) activity density, (h) absolute evenness, (i) inverse Simpson index, and (j) Shannon entropy. (k) Effect of crop configuration on ground beetle taxonomic richness for each combination of location, year and crop. Barley-mixture consists of a mixture of barley-bean (2020) or barley-pea (2021). Squares indicate estimated means, the bar indicates the 95% confidence interval. Asterisks indicate significant differences among the crop configurations. Empty panels indicate combinations of years and locations that were not sampled. When no estimated mean and confidence interval are shown, crops were not grown or sampled in that year. Open circles indicate individual year series to visualize sample size (Table S7).

Strip cropping enhances ground beetle activity density

Ground beetle activity density was on average 30% higher in strip cropping fields than in monoculture fields (β0 = 0.303, SE = 0.121, p = 0.012; Fig. 1e), based on rarefaction. However, there was no significant difference in activity density between the strip cropping fields and monocultures that harboured the richest beetle communities (β0 = -0.110, SE = 0.088, p = 0.215). Crop-level activity density of ground beetles was not affected by crop configuration (Fig. 1g, S1a,b).

Added in the revised version: Crop configuration alters abundance of abundant genera

Added in the revised version: We tested how crop configuration affected the abundance of twelve abundant ground beetle genera (Table S3). We analysed this separately per location as some genera only occurred at specific locations. Four genera were more abundant in strip cropped fields in at least one location (Anchomenus, Bembidion, Harpalus and Nebria), whereas four genera were more abundant in monocultures in at least one location (Amara, Calathus, Pterostichus and Trechus) (Fig 2). The other four common genera (Blemus, Clivina, Loricera and Poecilus) were not significantly influenced by crop configuration. Furthermore, no ground beetle genus showed significantly contrasting responses to crop configuration among locations.

Added in the revised version: Ground beetle species associated with crop configuration (monoculture (a) versus strip cropping (b)).

Added in the revised version: Results obtained by generalized linear mixed models on the twelve most common genera of ground beetles with the four locations analysed separately (Table S5). Only those genera are given for which cropping system significantly influence activity density within at least one location (α = 0.05). Locations were the genus had a higher activity density in one of the cropping systems are indicated between brackets (Al=Almere; Le = Lelystad; Va = Valthermond; Wa = Wageningen). For Pterostichus, data from Almere in 2020 was analysed separately, as models did not fit elsewise. The asterisk “*” next to Pterostichus indicates that only in 2020 the difference between monoculture and strip cropping was significant. Photo credit: Ortwin Bleich, retrieved from: www.eurocarabidae.de.

Crop configuration alters ground beetle community composition

Earlier version: Ground beetle communities were significantly influenced by crop configuration, but the effects were highly dependent on the specific context created by the combination of location, year and crop (Fig. S3-S8, Table 1, S3). The context-dependency of configuration effects on ground beetle communities is illustrated in redundancy analyses of ground beetle assemblages per crop combination at Wageningen in 2021 and 2022 (Fig. S9). Here, we found distinct ground beetle communities among crop configurations for pumpkin, barley and potato in 2021, and for cabbage and oat in 2022. In the other cases the difference between crop configurations was not significant (Fig. S9, Table S4). Moreover, in all crop combinations except for potato-grass in 2021, the difference in ground beetle communities between the monocultures of the constituent crops were significant, while this was never the case for ground beetle communities of crops in strip cropping (Table S4). This indicates that strip cropping might lead to overlapping crop-related communities. However, these results could be spatially autocorrelated as samples from different crops were in closer proximity of each other in strip cropping than among monocultures (Fig S10d).

Revised version: Ground beetle communities were significantly influenced by crop configuration, but the effects were highly dependent on the specific context created by the combination of location, year and crop (Fig. S3-S8, Table 1, S4). The context-dependency of configuration effects on ground beetle communities is illustrated in redundancy analyses of ground beetle assemblages per crop combination at Wageningen in 2021 and 2022 (Fig. S9). Here, we found distinct ground beetle communities among crop configurations for pumpkin, barley and potato in 2021, and for cabbage and oat in 2022. In the other cases the difference between crop configurations was not significant (Fig. S9, Table S5). Moreover, in all crop combinations except for potato-grass in 2021, the difference in ground beetle communities between the monocultures of the constituent crops were significant, while this was never the case for ground beetle communities of crops in strip cropping (Table S5). This indicates that strip cropping might lead to overlapping crop-related communities. However, these results could be spatially autocorrelated as samples from different crops were in closer proximity of each other in strip cropping than among monocultures (Fig S10d).

Removed in the revised version: Crop configuration alters abundance of common genera

Removed in the revised version: We tested how crop configuration affected twelve common ground beetle genera (Table S5). We analysed this separately per location as some genera only occurred at specific locations. Four genera were more abundant in strip cropped fields in at least 1 location (Anchomenus, Bembidion, Harpalus and Nebria), whereas four genera were more abundant in monocultures in at least 1 location (Amara, Calathus, Pterostichus and Trechus) (Fig 2). The other four common genera (Blemus, Clivina, Loricera and Poecilus) were not significantly influenced by crop configuration. Furthermore, no ground beetle genus showed significantly contrasting responses to crop configuration among locations.

Removed in the revised version: Ground beetle species associated with crop configuration (monoculture (a) versus strip cropping (b)).

Removed in the revised version: Results obtained by generalized linear mixed models on the 12 most common genera of ground beetles with the four locations analysed separately (Table S5). Only those genera are given for which cropping system significantly influence activity density within at least one location (α = 0.05). Locations were the genus had a higher activity density in one of the cropping systems are indicated between brackets (Al=Almere; Le = Lelystad; Va = Valthermond; Wa = Wageningen). For Pterostichus, data from Almere in 2020 was analysed separately, as models did not fit elsewise. The asterisk “*” next to Pterostichus indicates that only in 2020 the difference between monoculture and strip cropping was significant. Photo credit: Ortwin Bleich, retrieved from: www.eurocarabidae.de.

Crop configuration does not increase the number of rare ground beetle species

Earlier version: Among the 461 year series, we only found two rare species (according to waarneming.nl): one individual of Microlestes minutulus, and five individuals of Harpalus signaticornis. This latter species was found most often in a wheat monoculture, but this was likely due to a failed crop which created a very open habitat. Harpalus signaticornis is known to inhabit recently disturbed, dry, open habitats with limited vegetation (Turin et al., 2012), a situation similar to this sparsely-covered monoculture. All other species were common or relatively common species (Table S2). Furthermore, most ground beetle species were either ruderal habitat specialists or eurytopic species that occur in many different habitats (Table S2).

Revised version: Among the 461 year series, we only found two rare species (following waarneming.nl): one individual of Microlestes minutulus, and five individuals of Harpalus signaticornis. This latter species was found most often in a wheat monoculture, but this was likely due to a failed crop which created a very open habitat. Harpalus signaticornis is known to inhabit recently disturbed, dry, open habitats with limited vegetation (Turin et al., 2012), a situation similar to this sparsely-covered monoculture. All other species were common or relatively common species (Table S2). Furthermore, most ground beetle species were either ruderal habitat specialists or eurytopic species that occur in many different habitats (Table S2).

Discussion

It is well established that different crop types have distinct ground beetle communities (Eyre et al., 2009; Holland and Luff, 2000) and that increasing habitat diversity by including multiple crops in a field can enhance ground beetle diversity (Cuperus et al., 2024; Puliga et al., 2023). Our study shows that strip cropping increased field-level ground beetle richness by 15%. However, ground beetle communities within the same crop in a strip or in monoculture were mostly similar. This indicates that the 15% increase in richness at the field level can be mostly attributed by the higher number of crops in strip cropped fields that harboured crop-related ground beetle communities, and that there was only limited mixing of ground beetle species among crops in strip crops. This is in line with earlier findings that ground beetle movement is reduced by crop edges (Allema et al., 2014; Anderson et al., 2024). Further research on movement behaviours of ground beetles at crop edges might help explain how ground beetles distribute themselves within a strip cropping field, and whether they utilize the different resources provided by a more diverse cropping system.

The 30% higher ground beetle activity density in strip cropping fields compared to monoculture fields may be explained by a more stable and diverse habitat with refuges and alternative resources in strip crops (Ratnadass et al., 2012). Crop diversification enhances prey biomass for ground beetles (Lichtenberg et al., 2017) and increases weed seed richness as compared to monocultures (Ditzler et al., 2023), both of which reduce bottom-up control by increased food provision (Carbonne et al., 2022). Alternatively, the increase in activity density might be caused by higher movement of ground beetles, which in turn can be the consequence of food starvation (Wallin and Ekbom, 1994). Indeed, several papers show reduced abundances of herbivorous insects in strip- and intercropping (Alarcón Segura et al., 2022; Cuperus et al., 2023; Rakotomalala et al., 2023), which may be potential prey for ground beetles (Turin, 2000). Furthermore, strip cropping systems may support high herbivore reproduction in combination with a high predation pressure resulting in herbivore populations dominated by early life stages (Karssemeijer et al., 2024). Therefore, strip cropping may favour ground beetles that predate on, for instance, herbivore eggs, such as Bembidion spp. and Anchomenus dorsalis, which were more abundant in strip cropping (Finch, 1996), and disadvantaging ground beetles that feed on larger prey, such as Pterostichus melanarius, which was less abundant in strip cropping. Multi-taxa evaluation of future strip cropping studies will be a valuable approach to increase understanding of biomass flows and trophic interactions within diverse agricultural systems.

The 15% increase in ground beetle species richness through strip cropping is mostly caused by an increase in species (relatively) common to agricultural fields, rather than an increase in rare species or species otherwise common in other habitat types. A keystone species here might be Pterostichus melanarius, the most common Pterostichus species in our samples. Pterostichus melanarius can be very dominant in pitfall traps, as we regularly found hundreds of individuals per trap. This relatively large ground beetle might compete with or predate on other ground beetle species (Roubinet et al., 2018). Pterostichus melanarius is especially well-adapted to highly intensive agriculture (Turin, 2022, 2000), which explains why it was more abundant in monocultures previously (Ditzler et al., 2021). In turn, the reduced P. melanarius populations in strip cropping might allow other species to persist, such as Harpalus spp. and Anchomenus dorsalis. Therefore, strip cropping might facilitate ground beetle biodiversity by reducing the dominance of a common, competitive species.

Biodiversity gains in our study did in most cases not coincide with productivity losses (Table S6). Earlier studies on crop yield in the strip cropping fields in Lelystad and Wageningen during the years of pitfall trapping show a yield decrease when strip cropping cabbage (Carrillo-Reche et al., 2023) and wheat (Ditzler et al., 2023), whereas potato yield was unaffected by crop configuration (Ditzler et al., 2023). In Almere bean and parsnip yields were higher in strip cropping, whereas oat and onion yield were lower (Juventia and van Apeldoorn, 2024). In Valthermond crop productivity was similar for monocultures and strip cropping (Table S6). Therefore, biodiversity gains do not compromise agricultural production.

We show that changing crop configuration from monoculture to strip cropping, on average, enhances ground beetle richness by 15% and activity density by 30% within agricultural fields. These results show that strip cropping can lead to increases in biodiversity that approach those achieved by shifting from conventional to organic farming practices (+19% richness, Lichtenberg et al., 2017; +23% richness, Gong et al., 2022) and by other in-field diversification measures, like hedgerows and flower strips (+23% richness, Lichtenberg et al., 2017; +24% richness, Beillouin et al., 2021). While organic management or in-field diversification measures generally lead to lower productivity (Gong et al., 2022), strip cropping can maintain crop productivity without taking land out of production (Campanelli et al., 2023; Juventia and van Apeldoorn, 2024; Van Oort et al., 2020), although occasional yield reductions have been reported (Carrillo-Reche et al., 2023; Ditzler et al., 2023). The biodiversity gain in our study was achieved without additional crop-level diversification strategies designed for biodiversity conservation, such as cover cropping or (flowering) companion plants. This increase in biodiversity stacks on top of already higher biodiversity achieved through organic management (Gong et al., 2022; Lichtenberg et al., 2017). Furthermore, the estimated richness effects of strip cropping, i.e., no change at crop level and 15% increase at field level, are conservative because they only consider the pairwise comparison of biodiversity of two or three crops in a strip cropping configuration to monocultures, whereas the inclusion of more crops within strip cropping could further enhance ground beetle biodiversity. Also, the effect of strip cropping on ground beetle biodiversity might be more pronounced in conventional agriculture and at larger field sizes, as the relatively small-scale, organic fields that we studied might have already had a relatively high ground beetle biodiversity (Tscharntke et al., 2021). The ground beetle communities in our study were dominated by farmland and eurytopic species and contained only two rare species (Turin, 2022). Future research could test whether the inclusion of other in-field diversification measures within strip-cropped fields, such as the establishment of perennial semi-natural habitats (Rischen et al., 2021; Sirami et al., 2019) or uptake at larger spatial extents (Tscharntke et al., 2021) would allow ground beetles with other habitat preferences to establish in agricultural fields.

Materials and methods

Study area

A multi-location study was conducted on four organic farms across the Netherlands (Fig. 1a). Three experimental farms were managed by Wageningen University & Research (Lelystad, Valthermond, Wageningen) and one commercial farm was managed by Exploitatie Reservegronden Flevoland (ERF B.V.) located in Almere. All four locations contained both strip cropping and monocultural crop fields, but differed in soil type, establishment year of the strip cropping experiment, number of crops grown, length of the crop rotation, number of sampled crops and sampling years, and farm and landscape characteristics such as percentage of on-farm semi-natural habitat (SNH), mean field size, and landscape configuration (Fig. 1b, S10). The locations Almere and Lelystad were located in a homogeneous, open polder landscape characterized by intensive arable crop production and non-crop habitats consisting of grass margins, tree lines and watercourses. Valthermond was located in an open, reclaimed peat landscape with intensive arable crop production characterized by long and narrow fields separated by grassy margins and ditches and limited areas of woody elements. The site at Wageningen was located in a more complex landscape with smaller field sizes, and non-crop habitat consisting of woodlots, hedgerows, tree lines, ditches and farmyards.

Experimental lay-out

At Almere and Valthermond the crops in strip cropping were all grown alongside each other, whereas at Lelystad and Wageningen two alternating crops (crop pairs) were grown alongside each other (Fig. 1b, S10). At each location, strip cropping and monoculture fields were always paired on the same experimental field (Fig. S10). At Almere eight different crops were grown in alternating strips of 6 m width, including celeriac (Apium graveolens var. rapaceum), broccoli (Brassica oleracea var. italic), oat (Avena sativa), onion (Allium cepa), parsnip (Pastinaca sativa), faba bean (Vicia faba), potato (Solanum tuberosum) and a mix of ryegrass and white clover referred to as grass-clover (Lolium perenne/Trifolium repens) (Fig. S10a). At Lelystad four different crop pairs were grown in alternating strips of 3 m width, including carrot (Daucus carota subsp. sativus) and onion, white cabbage (Brassica oleracea var. capitata) and wheat (Triticum aestivum), sugar beet (Beta vulgaris) and barley (Hordeum vulgare), and potato and ryegrass (Fig. S10b). At Valthermond eight different crops were grown in alternating strips of 6 m width, including potato, barley, barley mixed with broad bean (Vicia faba) in 2020, barley mixed with pea (Pisum sativum) in 2021, sweet corn (Zea mays convar. saccharata var. rugosa), sugar beet, common bean (Phaseolus vulgaris), and grass-clover (Fig. S10c). At Wageningen three different crop pairs were grown in alternating strips of 3 m width, including white cabbage and wheat (2019 – 2021) or oat (2022), barley and pumpkin (Cucurbita maxima, 2020 - 2022) or bare soil due to crop failure (2019), and potato and ryegrass (Fig. S10d). The crop combinations and neighbors were selected based on literature, expert knowledge and experience of functionality in terms of expected advantages for yield and pest and disease control. Large scale monoculture plots (0.25 ha to 2.30 ha) served as reference, hereafter referred to as monoculture (Fig. S10). At Lelystad, Valthermond and Wageningen not each crop grown in strips was present as monoculture in each year, but only those crops for which a monoculture was present were sampled (Fig. 1k, Table S7). All fields were managed according to organic regulations, yet at each location fertilization and weed management reflected regional practices and were adjusted to local soil conditions. Flower strips were sown within the experimental fields at Almere and Valthermond (Fig. S10, Table S8).

Sampling

Earlier version: The ground beetle community was sampled using pitfall traps in all crops for which both a monoculture and strip cropping field were present, at each location and in multiple rounds per year between March and September. The sampled crops, number of rounds and number of pitfalls differed per year and per location (Table S7). Pitfall traps consisted of a transparent plastic cup (9.2 cm diameter, 14 cm height) placed in the soil so the top of the cup was level with the soil surface. We did not use funnels inside the cups. Pitfalls were filled with approximately 100 ml water mixed with non-perfumed soap (around 3 cm heigh water level) and covered with a black plastic rain guard (12.5 cm diameter) around 2 to 5 cm above the soil surface. Pitfall traps were placed in the center of a strip (1.5 – 3 meters from the edges of the strip), and at least 10 meters from field edges at fixed locations at different moments in a year. Furthermore, traps were mostly placed at equal distances from the field edges in strip-cropped and monocultural fields. At Almere and Valthermond, the number of pitfalls in monocultures and strip cropping was the same. At Wageningen and Lelystad, the number of pitfalls in strip cropped fields was usually lower than in the monoculture because the area of the strips was only half of that of the monocultures.

Revised version: The ground beetle community was sampled using pitfall traps in all crops for which both a monoculture and strip cropping field were present, at each location and in multiple rounds per year between March and September. The sampled crops, number of rounds and number of pitfalls differed per year and per location (Table S7). Pitfall traps consisted of a transparent plastic cup (9.2 cm diameter, 14 cm height) placed in the soil so the top of the cup was flush with the soil surface. We did not use funnels inside the cups. Pitfalls were filled with approximately 100 ml water mixed with non-perfumed soap (around 3 cm heigh water level) and covered with a black plastic rain guard (12.5 cm diameter) around 2 to 5 cm above the soil surface. Pitfall traps were placed in the center of a strip (1.5 – 3 meters from the edges of the strip), and at least 10 meters from field edges at fixed locations at different moments in a year. Furthermore, traps were mostly placed at equal distances from the field edges in strip-cropped and monocultural fields. At Almere and Valthermond, the number of pitfalls in monocultures and strip cropping was the same. At Wageningen and Lelystad, the number of pitfalls in strip cropped fields was usually lower than in the monoculture because the area of the strips was only half of that of the monocultures.

Ground beetle identification

Earlier version: Ground beetles were identified up to species level for Almere in 2021 and 2022, Lelystad and Wageningen. At Almere in 2020, identification was kept at the genus level due to an extremely high abundance of Pterostichus melanarius/niger and Poecilus cupreus/versicolor, and the associated high time investment for identification. At Valthermond, due to labour constraints, we chose to consider species that are complicated to distinguish only up to genus level (like Poecilus cupreus/versicolor, Harpalus spp. other than H. rufipes and H. affinis, Bembidion spp.). For all analyses, year series were made, in which all ground beetle catches from the same pitfall trap were pooled per year.

Revised version: Ground beetles were identified up to species level for Almere in 2021 and 2022, Lelystad and Wageningen. At Almere in 2020, identification was kept at the genus level due to an extremely high abundance of Pterostichus melanarius/niger and Poecilus cupreus/versicolor, and the associated high time investment for identification. At Valthermond, due to labour constraints, we chose to consider species that are complicated to distinguish only up to genus level (like Poecilus cupreus/versicolor, Harpalus spp. other than H. rufipes and H. affinis, Bembidion spp.). For all analyses, year series were made, in which all ground beetle catches from the same pitfall trap were pooled per year. We examined habitat preferences of ground beetle using Turin et al. (2022) and ground beetle rarity using www.waarneming.nl (checked on 15-01-2025).

Statistical analyses

We used R, version 4.2.2 for all statistical analyses.

Effect of crop configuration on field-level richness and activity density

Earlier version: To analyse the difference in species richness and activity density between monocultures and strip cropping configurations at field level, we used rarefaction of samples within the same field. To rarify to an equal sampling intensity, we calculated the average cumulative number of species or individuals within x year series, where x is the largest number of year series available for the crop configuration comparison (correcting for unequal sampling between monocultures and strip cropping, or for missing samples). Next, we calculated the relative change due to strip cropping by subtracting the number of species or individuals found in the monoculture field from the number found in the strip cropping field and then dividing the result by the number of species or individuals in the monoculture. This gave the relative change centered around zero, where negative values indicated higher richness in monocultures and positive values higher richness in strip cropping. We then analysed this data using Generalized Linear Mixed Models (GLMM) with a Gaussian distribution and assessed whether the intercept deviated significantly from zero. As random variables, we used location and year, with year nested in location. We ran these analyses using a dataset that included all comparisons among monocultural fields and strip cropping fields of all locations. However, as in this case the strip cropped field was compared with both monocultural fields of the corresponding crop pair, we also tested the effect of field-level richness using only one monoculture per strip cropped field. To obtain a conservative estimate of the effect of strip cropping, we chose the monocultural fields with the highest taxonomic richness or activity density among the constitutive crops of the strip cropping fields. Generalized linear models were run using the glmmTMB package (Brooks et al., 2017) and tested for model fit using the DHARMa package (Hartig, 2018).

Revised version: To analyse the difference in species richness and activity density between monocultures and strip cropping configurations at field level, we used rarefaction of samples within the same field. To rarify to an equal sampling intensity, we calculated the average cumulative number of species or individuals within x year series, where x is the largest number of year series available for the crop configuration comparison (correcting for unequal sampling between monocultures and strip cropping, or for missing samples). Next, we calculated the relative change due to strip cropping by subtracting the number of species or individuals found in the monoculture field from the number found in the strip cropping field and then dividing the result by the number of species or individuals in the monoculture (Zou et al., 2020). This gave the relative change centered around zero, where negative values indicated higher richness in monocultures and positive values higher richness in strip cropping. We then analysed this data using Generalized Linear Mixed Models (GLMM) with a Gaussian distribution and assessed whether the intercept deviated significantly from zero. As random variables, we used location and year, with year nested in location. We ran these analyses using a dataset that included all comparisons among monocultural fields and strip cropping fields of all locations. However, as in this case the strip cropped field was compared with both monocultural fields of the corresponding crop pair, we also tested the effect of field-level richness using only one monoculture per strip cropped field. To obtain a conservative estimate of the effect of strip cropping, we chose the monocultural fields with the highest taxonomic richness or activity density among the constitutive crops of the strip cropping fields. Generalized linear models were run using the glmmTMB package (Brooks et al., 2017) and tested for model fit using the DHARMa package (Hartig, 2018).

Effect of crop configuration on crop-level biodiversity

To quantify biodiversity we used five variables: (1) activity density, the total number of ground beetles found per year series; (2) taxonomic richness, the total number of species or genera (lowest taxonomic level available) found per year series; (3) the inverse Simpson index, the inverse of the sum of proportions of different species over the total abundance (Simpson, 1949); (4) absolute evenness, the number of effective species calculated by dividing the inverse Simpson index by the taxonomic richness (Williams, 1964); and (5) Shannon entropy (Shannon and Weaver, 1949). We chose absolute evenness as our measure for evenness, as this method removes the richness component from the inverse Simpson index and adheres to all requirements for an evenness index (Smith and Wilson, 1996; Tuomisto, 2012). We included both the inverse Simpson index and Shannon entropy as the former is more sensitive to changes in evenness and the latter to species richness (DeJong, 1975).

Earlier version: To analyse the effect of crop configuration on total ground beetle activity density, taxonomic richness, evenness, inverse Simpson index and Shannon entropy we used GLMM. We constructed models for each response variable, using data from all four locations. In these models we included crop configuration (monoculture or strip cropping) as a fixed factor. We included location, year and crop as nested random variables in these models, with crop nested in year and year in location. To quantify and visualize the variation in responses between locations, years and crops, we ran generalized linear models (GLM) with a variable that combined these three variables into one, which was also included as a fixed factor. Here, we also included the interaction between crop configuration and the combined variable for crop, location and year. For the model on activity density we used the negative binomial distribution (log link function), as this was count data; for richness, evenness and Shannon entropy we used Gaussian distribution; and for the inverse Simpson index we used a gamma distribution (inverse link function) as a Gaussian distribution did not fit well. We used the ‘DHARMa’ package to validate model assumptions. We used estimated marginal means to assess differences between monoculture and strip cropping.

Revised version: To analyse the effect of crop configuration on total ground beetle activity density, taxonomic richness, evenness, inverse Simpson index and Shannon entropy we used GLMM. We constructed models for each response variable, using data from all four locations. In these models we included crop configuration (monoculture or strip cropping) as a fixed factor. We included location, year and crop as nested random variables in these models, with crop nested in year and year in location. To quantify and visualize the variation in responses between locations, years and crops, we ran generalized linear models (GLM) with a variable that combined these three variables into one, which was also included as a fixed factor. Here, we also included the interaction between crop configuration and the combined variable for crop, location and year. For the model on activity density we used the negative binomial distribution (log link function), as this was count data; for richness, evenness and Shannon entropy we used Gaussian distribution; and for the inverse Simpson index we used a gamma distribution (inverse link function) as a Gaussian distribution did not fit well. We used the ‘DHARMa’ package to validate model assumptions (Hartig, 2018). We used estimated marginal means to assess differences between monoculture and strip cropping (Hothorn et al., 2016; Lenth, 2018).

To analyse the effect of crop configuration on the activity density of the twelve most common ground beetle genera, we used another GLMM with negative binomial distribution. Fixed effects were crop configuration, location and their interactions, and crop nested in year was added as a random effect. As certain genera only occurred at specific locations, we only included locations where the genus was found. All models were tested for model fit using the DHARMa package. We used estimated marginal means to assess differences between monoculture and strip cropping.

Community composition of crop configurations and crops

To assess whether crop configurations have distinct ground beetle communities we used permanova with a Hellinger transformation (with 999 permutations) using the “vegan” package (Oksanen et al., 2013). We used a Hellinger transformation to give more weight to rarer species, thus accounting for highly abundant species in some samples (such as P. melanarius and H. rufipes). We only included species that occurred in at least 3% of the pitfall samples to avoid strong influence of very rare species. We used data from locations and years where pitfall catches had been identified to species level (Almere 2021/22, Lelystad and Wageningen). We considered four models, one for all locations combined and one for each location separately. In all models we included crop configuration and the nested variables of location (whenever applicable), year and crop as fixed factors. We also analysed the interaction between crop configuration and this nested structure of location, year and crop. To visualize these results, we used non-metric multidimensional scaling (NMDS) on Hellinger transformed data. To show whether crop configurations have distinct ground beetle communities, we used redundancy analysis (RDA) on Hellinger transformed data. Here, we again conducted four analyses, one for all locations combined and one for each of the three considered locations. We only used crop configuration as a predictor to force RDA to show any change in ground beetle community associated with crop configuration. As such, only one RDA axis was created per model, which was plotted against the first principal component describing the residual variation.

Due to the large influence of location and year on ground beetle communities, visualizing any general effects of strip cropping on these communities using all available data was challenging. To address this, we conducted RDA and visualized the effect of crop configuration on a subset of the data from one location and one year. We chose the data from Wageningen in 2021 and 2022 for this analysis because it provided a set-up where strip cropping of two crops could be compared with their constituent monocultures within the same experimental fields. Here, we used both crop configuration, crop and their interaction term as explanatory variables for each crop pair separately. To analyse whether ground beetle communities significantly differed among combinations of crops and crop configurations we ran pairwise permanova on all three fields and two years separately using the “pairwiseAdonis” package (Martinez Arbizu, 2020).